Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Aaron Holz"
COLLABORATORS = ""
```

---

---
## Aufgabe 1 (6 Punkte)

Nenne 6 Teilbereiche der Physik!


<!-- grade_id: cell-595ee723f75ca4ab, points: 6, grade: true, solution: true -->
Mechanik, Akustik, Magnetismus, Elektrizität, Optik, Atomphysik

---
## Aufgabe 2 (3 Punkte)

Gib drei Beispiele an, die zeigen, dass die Erkenntnisse der Physik für den Menschen nutzbar gemacht werden können.


<!-- grade_id: cell-174a6c14d7db1673, points: 3, grade: true, solution: true -->
1) Die Benützung einer Lupe, um etwas größer zu sehen.

2) Die Erzeugung von elektrischem Strom in Wasserkraftwerken.

3) Die Benützung eines Kompasses zur Orientierung.

---
## Aufgabe 3 (2 Punkte)
Ergänze: Der geografische
Nordpol der Erde ist ein
magnetischer ...........
der geografische Südpol der
Erde ist ein
magnetischer ...............


<!-- grade_id: cell-0e609bc0dfa149a2, points: 2, grade: true, solution: true -->
Der geografische Nordpol der Erde ist ein magnetischer
Südpol.

der geografische Südpol der Erde ist ein magnetischer
Nordpol.

## Aufgabe 4 (3 Punkte)

Woraus besteht jedes Atom?

<!-- grade_id: cell-0aba32582328c6dd, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Aufgabe 5 (2 Punkte)

Ein Körper bewegt sich gleichmäßig beschleunigt und erreicht in der Zeit 12,0s eine Geschwindigkeit von 72,0 $\frac{m}{s}$. Berechne die Beschleunigung a des Körpers.

<!-- grade_id: cell-7d67a47726dd46ae, points: 2, grade: true, solution: true -->
$a = \frac{v}{t} \Rightarrow a = \frac{{72,0\frac{{\rm{m}}}{{\rm{s}}}}}{{12,0{\rm{s}}}} = 6,00\frac{{\rm{m}}}{{{{\rm{s}}^{\rm{2}}}}}$

## Bonusaufgabe 6 (+1 Punkte)

Erstelle ein Python-Programm, das die Beschleunigung des Körpers von Zeit und Geschwindigkeit berechnen kann.

<!-- grade_id: cell-e32eab35f8b5279a, points: 0, grade: false, solution: true -->
```python
def beschleunigung(zeit, geschwindigkeit):
    ### BEGIN SOLUTION
    if zeit < 1:
        raise ValueError("Die Zeit kann nicht negativ sein")
    return geschwindigkeit/zeit
    ### END SOLUTION
```

<!-- grade_id: cell-c08ada6e138ae0da, points: 1, grade: true, solution: false -->
```python
"""Tests"""
assert beschleunigung(10,20) == 2.0

### BEGIN HIDDEN TESTS 
assert beschleunigung(10,-20) == -2.0
### END HIDDEN TESTS
```